[1613641002] SERVICE ALERT: localhost;Root Partition;WARNING;SOFT;1;DISK WARNING - free space: / 14304 MiB (19.84% inode=65%):
[1613641062] SERVICE ALERT: localhost;Root Partition;WARNING;SOFT;2;DISK WARNING - free space: / 13710 MiB (19.02% inode=64%):
[1613641182] SERVICE ALERT: localhost;Root Partition;WARNING;SOFT;3;DISK WARNING - free space: / 12958 MiB (17.98% inode=64%):
[1613641242] SERVICE ALERT: localhost;Root Partition;WARNING;HARD;4;DISK WARNING - free space: / 12134 MiB (16.83% inode=63%):
[1613641338] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;1;WARNING - load average: 6.56, 4.17, 2.63
[1613641416] SERVICE ALERT: localhost;Current Load;OK;SOFT;2;OK - load average: 3.20, 3.62, 2.54
[1613642676] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;1;WARNING - load average: 7.56, 3.69, 2.54
[1613642737] SERVICE ALERT: localhost;Current Load;OK;SOFT;2;OK - load average: 3.83, 3.40, 2.52
[1613643397] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;1;WARNING - load average: 5.14, 3.27, 2.71
[1613643457] SERVICE ALERT: localhost;Current Load;OK;SOFT;2;OK - load average: 3.19, 3.07, 2.69
[1613645617] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 16.36, 7.43, 4.51
[1613645668] SERVICE ALERT: viewnext105.sinensia.pw;gitlab;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1613645678] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 86.70, 28.16, 11.75
[1613645738] SERVICE ALERT: viewnext105.sinensia.pw;gitlab;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1613645739] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 123.94, 49.04, 19.93
[1613645800] SERVICE ALERT: localhost;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 148.85, 69.40, 28.73
[1613645811] SERVICE ALERT: viewnext105.sinensia.pw;gitlab;CRITICAL;SOFT;3;CRITICAL - Socket timeout
[1613645882] SERVICE ALERT: viewnext105.sinensia.pw;gitlab;CRITICAL;HARD;4;CRITICAL - Socket timeout
[1613645888] SERVICE ALERT: viewnext105.sinensia.pw;jenkins;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1613645959] SERVICE ALERT: viewnext105.sinensia.pw;jenkins;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1613646030] SERVICE ALERT: viewnext105.sinensia.pw;jenkins;CRITICAL;SOFT;3;CRITICAL - Socket timeout
[1613646091] SERVICE ALERT: viewnext105.sinensia.pw;jenkins;CRITICAL;HARD;4;connect to address 10.250.5.20 and port 9080: Connection refused