[1613556926] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613556986] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613557039] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 22: Connection refused
[1613557046] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613557076] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613557099] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 22: Connection refused
[1613557106] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613557136] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613557159] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 22: Connection refused
[1613557196] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613557219] SERVICE ALERT: localhost;SSH;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 22: Connection refused
[1613557256] SERVICE ALERT: localhost;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613567107] SERVICE ALERT: viewnext105.sinensia.pw;jenkins;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 1529 bytes in 0.002 second response time
[1613567237] SERVICE ALERT: viewnext105.sinensia.pw;jenkins;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 1530 bytes in 0.002 second response time
[1613567297] SERVICE ALERT: viewnext105.sinensia.pw;jenkins;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 1531 bytes in 0.002 second response time
[1613567357] SERVICE ALERT: viewnext105.sinensia.pw;jenkins;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 1532 bytes in 0.002 second response time
[1613567608] SERVICE ALERT: localhost;HTTP;OK;HARD;4;HTTP OK: HTTP/1.1 200 OK - 886 bytes in 0.003 second response time
[1613567720] SERVICE ALERT: localhost;SSH;OK;HARD;4;SSH OK - OpenSSH_7.6p1 Ubuntu-4ubuntu0.3 (protocol 2.0)
[1613573877] SERVICE ALERT: viewnext105.sinensia.pw;jenkins;OK;HARD;4;HTTP OK: HTTP/1.1 200 OK - 46646 bytes in 0.279 second response time
[1613580558] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;1;WARNING - load average: 5.57, 2.72, 1.52
[1613580618] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;2;WARNING - load average: 7.99, 3.99, 2.04
[1613580678] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;3;WARNING - load average: 6.76, 4.33, 2.28
[1613580738] SERVICE ALERT: localhost;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 12.24, 6.11, 3.01
[1613581038] SERVICE ALERT: localhost;Current Load;WARNING;HARD;4;WARNING - load average: 2.33, 4.13, 3.05
[1613581338] SERVICE ALERT: localhost;Current Load;OK;HARD;4;OK - load average: 0.64, 1.95, 2.40